// intro2.rs
// Make the code print a greeting to the world
// by completing the println statement.
// Execute `rustlings hint intro2` for a hint.

// I AM NOT DONE

fn main() {
    println!("",
        greeting="Welcome to",
        adjective="the extraordinary",
        noun="Rustlings Polkadot Academy self-assessment!"
    );
}