// variables2.rs
// Make me compile by printing "42" from the else statement.

// I AM NOT DONE

fn main() {

    if x == 10 {
        println!("Ten!");
    } else {
    
    // Print out "42" from the provided string_from_number
    let string_from_number = String::from("0000420000");

    // TODO write println statement using string slice

    }
}
